---
title       : Chapter 4
description : Setting up projects, directories and installing packages



--- type:NormalExercise xp:100 skills:1 key:4303b8ac87
! [Alt text] (https://iu.box.com/s/7631mxi7eyczxfu41l51hqzgi7fhubol)
*** =instructions


*** =hint


*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Awesome!")
```