---
title       : Day 2
description : Data manipulation
attachments:
  slides_link: http://link.to.slides/
  
--- type:NormalExercise xp:100 skills:1 key:960f76760f


*** =instructions


*** =hint


*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Awesome!")
```