---
title       : Chapter 4
description : Setting up projects, directories and installing packages



--- type:NormalExercise xp:100 skills:1 key:4303b8ac87
## RStudio Review
! [Alt text] (https://iu.box.com/s/7631mxi7eyczxfu41l51hqzgi7fhubol)

Click on the link to view the screenshot

*** =instructions
- Examine the screenshot of RStudio
- In the R-Editor, complete the decsription of each of the four windowns in RStudio
- Choices are **console**, **packages, plots, help**, **script**, **history** 
*** =hint


*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
topleft <- ""

topright <- ""

bottomleft <- ""

bottomright <- ""
```

*** =solution
```{r}
topleft <- "script"

topright <- "history"

bottomleft <- "console"

bottomright <- "packages, plots, help"
```

*** =sct
```{r}
test_output_contains(`topleft <- "script"`, incorrect_msg = "try again.")
test_output_contains(`topright <- "history"`, incorrect_msg = "make sure your string is spelled correctly.")
test_output_contains(`bottomleft <- "console"`, incorrect_msg = "try again.")
test_output_contains(`bottomright <- "packages, plots, help"`, incorrect_msg = "try again.")
success_msg("Awesome!")
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:d729bc62a5
## Getting Help


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
