---
title       : Chapter 4
description : Setting up projects, directories and installing packages


--- type:MultipleChoiceExercise xp:50 skills:1 key:601d67379a
## RStudio Structure
![]()



--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:2035d5aaec
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
--- type:NormalExercise xp:100 skills:1 key:4303b8ac87

*** =instructions


*** =hint


*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Awesome!")
```