---
title       : Day 5
description : Text mining, text analysis
attachments:
  slides_link: http://link.to.slides/
--- type:NormalExercise xp:100 skills:1 key:1c7fe42bfe


*** =instructions


*** =hint


*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Awesome!")
```