---
title       : Chapter 4
description : Setting up projects, directories and installing packages



--- type:VideoExercise lang:r xp:50 skills:1 key:207105d363
## <<<New Exercise>>>


*** =video_link
https://iu.box.com/s/7631mxi7eyczxfu41l51hqzgi7fhubol

--- type:VideoExercise lang:r xp:50 skills:1 key:031a25d1d4
## <<<New Exercise>>>


*** =video_link
//player.vimeo.com/video/154783078

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:2035d5aaec
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
--- type:NormalExercise xp:100 skills:1 key:4303b8ac87

*** =instructions


*** =hint


*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Awesome!")
```