---
title       : Day 3
description : Staistical analysis
attachments:
  slides_link: http://link.to.slides/
--- type:NormalExercise xp:100 skills:1 key:8156a925e3


*** =instructions


*** =hint


*** =pre_exercise_code


*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}
success_msg("Awesome!")
```